# 大学入試とAI

* 「[東ロボくん - Wikipedia](https://ja.wikipedia.org/wiki/%E6%9D%B1%E3%83%AD%E3%83%9C%E3%81%8F%E3%82%93)」については別の機会に・・・

## DNCL

* [独立行政法人 大学入試センター](https://www.dnc.ac.jp/)
    * [共通テスト手順記述標準言語 (DNCL) の説明](https://www.dnc.ac.jp/albums/abm.php?d=67&f=abm00000819.pdf&n=R4_%E5%85%B1%E9%80%9A%E3%83%86%E3%82%B9%E3%83%88%E6%89%8B%E9%A0%86%E8%A8%98%E8%BF%B0%E6%A8%99%E6%BA%96%E8%A8%80%E8%AA%9E%EF%BC%88DNCL%EF%BC%89%E3%81%AE%E8%AA%AC%E6%98%8E.pdf)
    * [令和７年度以降の試験に向けた検討について | 独立行政法人 大学入試センター](https://www.dnc.ac.jp/kyotsu/shiken_jouhou/r7ikou/)
        * [平成 30 年告示高等学校学習指導要領に対応した令和７年度大学入学共通テストからの出題教科・科目 情報 サンプル問題](https://www.dnc.ac.jp/albums/abm.php?d=33&f=abm00000307.pdf&n=%E3%82%B5%E3%83%B3%E3%83%97%E3%83%AB%E5%95%8F%E9%A1%8C%E3%80%8E%E6%83%85%E5%A0%B1%E3%80%8F_%E5%95%8F%E9%A1%8C.pdf)

```python
%%ai_ask
DNCLについて
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
日本の大学入試で使われる標準言語の方
```

Output:
```
<IPython.core.display.Markdown object>
```

## R7: 情報 サンプル問題 第２問

### 問1

![fig-2-01.png](images/fig-2-01.png)

![fig-2-02.png](images/fig-2-02.png)

## 次のChatGPTとの応答例を参考に上の問題に解答せよ

```python
%%ai_ask
日本の比例代表選挙の基準得票数について
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
pythonでプログラムします。今回は組み込み関数だけ使います。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
配列Tomeiに "A党", "B党", "C党", "D党" を格納します。Tomeiは不変です。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
ちなみにpythonにおける不変な配列について
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
in english
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
Tomeiの最初の要素、最後の要素の取り出し方を説明してください。また、任意の場所iの要素の取り出し方も教えてください。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
次に、配列Tokuhyoに 1200, 660, 1440, 180 を格納します。Tokuhyoも不変です。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
変数sousuuを0で初期化し、Tokuhyoの総計を求めてsousuuに格納するプログラムを書いてください。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
変数sousuuを0で初期化し、Tokuhyoの総計を求めてsousuuに格納するプログラムを書いてください。
その際、for文を使ってループ変数を添字としてTokuhyoの各要素をアクセスしてください。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
参考までに、組み込み関数を使うと簡略化できますか？
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
次に、変数gisekiを導入し6で初期化してください。gisekiは定数です。変数kizyunsuuにsousuuをgisekiで割った値を格納してください。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
%%ai_continue
最後に、Tokuhyoの各要素をkizyunsuuで割って、Tomeiの値をラベルにして表示してください。
```

Output:
```
<IPython.core.display.Markdown object>
```

```python
Tokuhyo = (1200, 660, 1440, 180)
sousuu = sum(Tokuhyo)
giseki = 6
kizyunsuu = sousuu / giseki
Tomei = ("A党", "B党", "C党", "D党")

for i in range(len(Tokuhyo)):
    ratio = Tokuhyo[i] / kizyunsuu
    label = Tomei[i]
    print(f"{label}: {ratio:.2f}")
```

Output:
```
A党: 2.07
B党: 1.14
C党: 2.48
D党: 0.31
```

```python

```

---
## 回答

次のセルの "ア", "イ", "ウ" の値が ```"_"``` になっています。 
ダブルクオーテーション ```"``` で囲まれたアンダースコア ```_``` を回答 (半角英数字一文字) に書き換えて、
ツールバーの "Validate" を押してください。
ブルートフォースアタックができないように、応答が出るまで時間が掛かります。
すべて正答であれば "Success! Your notebook passes all the tests." と表示されます。

```python
your_answer = {
  "ア": "_",
  "イ": "_",
  "ウ": "_",
}
```

```python
%%ai_ask
ブルートフォースアタックとは
```

Output:
```
<IPython.core.display.Markdown object>
```

---
以下、採点用のセルにつき編集できない

```python
import hashlib
validate = {}
validate["ア"] = hashlib.md5(your_answer["ア"].encode('utf-8')).hexdigest()
validate["イ"] = hashlib.md5(your_answer["イ"].encode('utf-8')).hexdigest()
validate["ウ"] = hashlib.md5(your_answer["ウ"].encode('utf-8')).hexdigest()
```

<!-- grade_id: cell-0b6f370a8adb5128, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
print(hashlib.md5("3".encode('utf-8')).hexdigest())
print(hashlib.md5("8".encode('utf-8')).hexdigest())
print(hashlib.md5("b".encode('utf-8')).hexdigest())
### END HIDDEN TESTS
```

Output:
```
eccbc87e4b5ce2fe28308fd9f2a7baf3
c9f0f895fb98ab9159f51fd0297e236d
92eb5ffee6ae2fec3ad71c777531578f
```

<!-- grade_id: cell-78f3cbf5f3a78f54, points: 1, grade: true, solution: false -->
```python
assert validate["ア"] == "eccbc87e4b5ce2fe28308fd9f2a7baf3"
assert validate["イ"] == "c9f0f895fb98ab9159f51fd0297e236d"
assert validate["ウ"] == "92eb5ffee6ae2fec3ad71c777531578f"
```

```python

```